Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abdul Wahid"
COLLABORATORS = "2202396"
```

---

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function for Pythagorean theorem, or Pythagoras' theorem. Pythagorean is a fundamental relation in Euclidean geometry among the three sides of a right triangle. It states that the area of the square whose side is the hypotenuse (the side opposite the right angle) is equal to the sum of the areas of the squares on the other two sides. This theorem can be written as an equation relating the lengths of the legs a, b and the hypotenuse c, often called the Pythagorean equation:

$a^2 + b^2 = c^2$

Please write a function (code)to calculate the length of an unknown side of a right triangle, given two sides are known.

<!-- grade_id: myPythagorean_ans, points: 0, grade: false, solution: true -->
```python
import math
def myPythagorean(a,b):
    """write a function (code)to calculate the length of an unknown side 
    of a right triangle, given two sides are known.

    
    """
    import math
    c = (a*a)+(b*b)
    root = math.sqrt(c)
    print('The third side of square is = ', root)
    #return root
    raise NotImplementedError()
```

```python
myPythagorean(10,14)
```

Output:
```
The third side of square is =  17.204650534085253
17.204650534085253
```

<!-- grade_id: myPythagorean_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert myPythagorean(3,2) == 3.605551275463989
assert myPythagorean(7,3) == 7.615773105863909
```

Output:
```
The third side of square is =  3.605551275463989
The third side of square is =  7.615773105863909
```

---

## Question 2:

A `trapezoid` is a quadrilateral with two sides parallel. The `trapezoid` is equivalent to the British definition of the `trapezium`. An isosceles `trapezoid` is a `trapezoid` in which the base angles are equal so

<img src="https://www.w3resource.com/w3r_images/python-math-image-exercise-3.svg"
     alt="Markdown Monster icon"
     style="float: left; margin-right: 10px;" />
     
Please write a function (code)to calculate the area of trapezoid.

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def areaofTrap(base1, base2, height):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    import math
    A = 0.5*(base1+base2)*height
    print('The Area ofTrapezoid is = ', A)
    return A
    raise NotImplementedError()
```

```python
areaofTrap(5, 3, 6)
```

Output:
```
The Area ofTrapezoid is =  24.0
24.0
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert areaofTrap(10, 10, 20) == 200.0
assert areaofTrap(5, 10, 20) == 150.0
assert areaofTrap(5, 3, 5) == 20.0
```

Output:
```
The Area ofTrapezoid is =  200.0
The Area ofTrapezoid is =  150.0
The Area ofTrapezoid is =  20.0
```

---
## Question 3: 

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string(str1):
    """Write your code here to reverse the string"""
    
    # YOUR CODE HERE
    return str1[:: -1]
    raise NotImplementedError()
```

```python
reverse_string('CE880')
```

Output:
```
'088EC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string('Colchester') == 'retsehcloC'
assert reverse_string('University') == 'ytisrevinU'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```

## Question 4:
Write a Python program to print the odd numbers from a given list and in ascending order.

Sample List : `[1, 2, 3, 4, 5, 6, 7, 8, 9]` Expected Result : `[1, 3, 5, 7, 9]`

<!-- grade_id: odd_num_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def odd_num(lst_1):
    """Write a Python program to print the odd numbers from a given list and in ascending order."""
    
    # YOUR CODE HERE
    list_Odd = []
    for x in lst_1:
      if((x%2) !=0):
        list_Odd.append(x)
    list_Odd = sorted(list_Odd)
    return (list_Odd)
    raise NotImplementedError()
```

```python
odd_num([2, 3, 14, 17, 8, 19, 11, 14])
```

Output:
```
[3, 11, 17, 19]
```

<!-- grade_id: odd_num_test, points: 1, grade: true, solution: false -->
```python
assert odd_num([1, 2, 3, 4, 6, 7, 8, 9]) == [1, 3, 7, 9]
assert odd_num([2, 3, 4, 7, 8, 9, 11, 14]) == [3, 7, 9, 11]
assert odd_num([2, 3, 14, 17, 8, 19, 11, 14]) == [ 3, 11, 17, 19]
```

## Question 5:
Write a Python function to create and print a list where the values are square of numbers between `1` and `n`, where `n` is user defined number: Example `sqr_values(4) == [1, 4, 9]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def sqr_values(n):
    """Write a Python program to print the square of numbers between 1 and `n`"""
    
    # YOUR CODE HERE
    lst = []
    for y in range(1, n):
      lst.append(y*y)
    return(lst)
    raise NotImplementedError()
```

```python
sqr_values(11)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that sqr_values returns the correct answer for various inputs."""
assert sqr_values(4) == [1, 4, 9]
assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49]
assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

## Question 6: 
Write a Python program to generate all combinations of string with r length subsequences of elements from the input iterable.

A combination is a mathematical technique that determines the number of possible arrangements in a collection of items where the order of the selection does not matter. In combinations, you can select the items in any order.

Examples: comb_list('ABCD', 2) = [('A', 'B'), ('A', 'C'), ('A', 'D'), ('B', 'C'), ('B', 'D'), ('C', 'D')]

<!-- grade_id: comb_ans, points: 0, grade: false, solution: true -->
```python
import itertools as it
def comb_list(str_inp,r):
    # YOUR CODE HERE
    lst = []
   # from itertools import permutations
    combos = it.combinations(str_inp,r)
    for x in combos:
      lst.append(x)
    return lst
    raise NotImplementedError()
```

```python
comb_list('ABC',2)
```

Output:
```
[('A', 'B'), ('A', 'C'), ('B', 'C')]
```

<!-- grade_id: comb_test, points: 1, grade: true, solution: false -->
```python
assert comb_list('ABC',2) == [('A', 'B'), ('A', 'C'), ('B', 'C')]
assert comb_list('MNOP',3) == [('M', 'N', 'O'), ('M', 'N', 'P'), ('M', 'O', 'P'), ('N', 'O', 'P')]
```

```python

```